#### Exercise 1 - Binary Cross Entropy and its derivative
##### Task-1: Implement BCE loss function
##### Task-2: Implement the derivative of BCE loss function

### Import required libraries

```python
import numpy as np
```

### Task-1: Implement BCE loss function

<!-- grade_id: cell-f43878c7bcb7556f, points: 0, grade: false, solution: true -->
```python
def binary_cross_entropy(y_pred, y_true):

    return -(y_true * np.log(y_pred) + (1 - y_true) * np.log(1 - y_pred))
```

<!-- grade_id: cell-d509e8f5bbc0a548, points: 2, grade: true, solution: false -->
```python
# Test case: Imperfect prediction with floating-point values (y_pred = 0.8, y_true = 1)
y_pred, y_true = 0.8, 1
expected_output = -np.log(0.8)
output = binary_cross_entropy(y_pred, y_true)
assert np.isclose(output, expected_output, atol=1e-3), f"Expected {expected_output}, but got {output}"
```

### Task-2: Implement the derivative of BCE loss function

<!-- grade_id: cell-56b16bb374625825, points: 0, grade: false, solution: true -->
```python
def bce_derivative(y_pred, y_true):
   
    return (y_pred - y_true) / (y_pred * (1 - y_pred))
```

<!-- grade_id: cell-01c8b0a1f8bd3bb0, points: 2, grade: true, solution: false -->
```python
y_pred, y_true = 0.5, 1
expected_output = (0.5 - 1) / (0.5 * (1 - 0.5))  # Expected result should be -2.0
output = bce_derivative(y_pred, y_true)
assert np.isclose(output, expected_output, atol=1e-3), f"Expected {expected_output}, but got {output}"
```

```python
# binary classification with target label y = 1
y_true = 1
# Initial prediction probability
y_pred = 0.0085
```

```python
loss = binary_cross_entropy(y_pred, y_true)
print(f" y_true = {y_true}   y_pred = {y_pred}   Loss = {loss:0.2f}")
```

Output:
```
y_true = 1   y_pred = 0.0085   Loss = 4.77
```

```python
# binary classification with target label y = 0
y_true = 0
# Initial prediction probability
y_pred = 0.95
```

```python
loss = binary_cross_entropy(y_pred, y_true)
print(f" y_true = {y_true}   y_pred = {y_pred}   Loss = {loss:0.2f}")
```

Output:
```
y_true = 0   y_pred = 0.95   Loss = 3.00
```

```python
# exit(0)
```