#### Exercise 1 - Gradient Descent: Training and optimization
##### Task-1: Compute the gradient of quadratic loss function
##### Task-2: Train and update parameter to minimize the loss
##### Task-3: Compute the partial derivatives of quadratic loss function
##### Task-4: Train and update two parameters to minimize the loss

### Import required libraries

```python
import numpy as np
```

```python
def loss_fn(x):
    return x ** 2
```

### Task-1: Compute the derivative of loss_fun(x)

<!-- grade_id: cell-54324e617e6f4294, points: 0, grade: false, solution: true -->
```python
def loss_fn_derivative(x):
  
    return 2 * x
```

<!-- grade_id: cell-d9326a95ba3c34ba, points: 1, grade: true, solution: false -->
```python
x = 3
expected_output = 6
assert loss_fn_derivative(x) == expected_output, f"Expected {expected_output}, got {loss_fn_derivative(x)}"
```

```python
learning_rate = 0.001   # try with 0.01, 0.001, 1
```

```python
x = 8
```

### Task-2: Update parameter using the gradient and learning rate

<!-- grade_id: cell-f65c025420f1d30e, points: 0, grade: false, solution: true -->
```python
for _ in range(10):
    # Update the variable x by moving it in the direction that minimizes the loss. 
    # Use the learning rate and the derivative of the loss function to adjust x. 
    
    x = x - learning_rate * loss_fn_derivative(x)
    
    loss = loss_fn(x)
    print(f" x = {x:0.2f}   Loss = {loss:0.2f}")
```

Output:
```
x = 7.98   Loss = 63.74
 x = 7.97   Loss = 63.49
 x = 7.95   Loss = 63.24
 x = 7.94   Loss = 62.98
 x = 7.92   Loss = 62.73
 x = 7.90   Loss = 62.48
 x = 7.89   Loss = 62.23
 x = 7.87   Loss = 61.98
 x = 7.86   Loss = 61.73
 x = 7.84   Loss = 61.49
```

```python
def loss_fn2(x1, x2):
    return (x1-x2) ** 2
```

### Task-3: Return the partial derivatives of loss_fn2(x1, x2))

<!-- grade_id: cell-e1656332eded4f82, points: 0, grade: false, solution: true -->
```python
def loss_fn2_derivative(x1, x2):
    grad_x1 = 2*(x1-x2)
    # Compute partial derivative w.r.t x2
    grad_x2 = 2*(x1-x2) * -1

    return np.array([grad_x1, grad_x2])
```

<!-- grade_id: cell-2d3f83b012d9a04d, points: 1, grade: true, solution: false -->
```python
# Test case
x1_t, x2_t = 5, 3
expected_grad_x1 = 4
expected_grad_x2 = -4
grad = loss_fn2_derivative(x1_t, x2_t)
assert grad[0] == expected_grad_x1, f"Expected grad_x1 to be {expected_grad_x1}, but got {grad[0]}"
assert grad[1] == expected_grad_x2, f"Expected grad_x2 to be {expected_grad_x2}, but got {grad[1]}"
```

```python
x1 = 8
x2 = 4
```

```python
loss_fn2_derivative(x1, x2)
```

Output:
```
array([ 8, -8])
```

### Task-4: Update parameters using the partial derivative(s) and learning rate

<!-- grade_id: cell-986d3cbf838711f2, points: 0, grade: false, solution: true -->
```python
for _ in range(20):
    x1 = x1 - learning_rate * loss_fn2_derivative(x1, x2)[0]
    # Update parameter x2, use partial derivative w.r.t x2 and learning_rate
    x2 = x2 - learning_rate * loss_fn2_derivative(x1, x2)[1]
   
    loss1 = loss_fn2(x1, x2)
    print(f" x1 = {x1:0.2f}, x2 = {x2:0.2f}     Loss = {loss1:0.2f}")
```

Output:
```
x1 = 7.99, x2 = 4.01     Loss = 15.87
 x1 = 7.98, x2 = 4.02     Loss = 15.75
 x1 = 7.98, x2 = 4.02     Loss = 15.62
 x1 = 7.97, x2 = 4.03     Loss = 15.50
 x1 = 7.96, x2 = 4.04     Loss = 15.37
 x1 = 7.95, x2 = 4.05     Loss = 15.25
 x1 = 7.94, x2 = 4.06     Loss = 15.13
 x1 = 7.94, x2 = 4.06     Loss = 15.01
 x1 = 7.93, x2 = 4.07     Loss = 14.89
 x1 = 7.92, x2 = 4.08     Loss = 14.77
 x1 = 7.91, x2 = 4.09     Loss = 14.65
 x1 = 7.91, x2 = 4.09     Loss = 14.53
 x1 = 7.90, x2 = 4.10     Loss = 14.42
 x1 = 7.89, x2 = 4.11     Loss = 14.30
 x1 = 7.88, x2 = 4.12     Loss = 14.19
 x1 = 7.88, x2 = 4.12     Loss = 14.08
 x1 = 7.87, x2 = 4.13     Loss = 13.96
 x1 = 7.86, x2 = 4.14     Loss = 13.85
 x1 = 7.85, x2 = 4.15     Loss = 13.74
 x1 = 7.85, x2 = 4.15     Loss = 13.63
```