## Exercise 1 - Implement a simple math expression in Python
### Task-1: Simple Linear Function
### Task-2: Compute Output Values for Input Array
### Task-3: Linear Combination for Prediction
### Task-4: e power Function
### Task-5: log Function

```python
import numpy as np
```

## Task-1: Simple Linear Function

### Description: Implement a function f(x) that takes an input x and returns the value of the linear equation 2x + 1.
#### Input: An integer or a float x.
#### Output: The result of the expression 2x + 1.

<!-- grade_id: cell-3acd16b1c1fa9996, points: 0, grade: false, solution: true -->
```python
def f(x):
    return 2*x + 1
```

```python
x = 2.0
f(x)
```

Output:
```
5.0
```

<!-- grade_id: cell-a43b340e00800983, points: 2, grade: true, solution: false -->
```python
assert f(2.0) == 5.0
```

```python
x_values = np.array([-3, -1, 5, 2])
x_values
```

Output:
```
array([-3, -1,  5,  2])
```

## Task-2: Compute Output Values for Input Array

### Description: Given an above x_values, create y_values where each element is the output of the function f(x) applied to the corresponding element of x_values.

#### Result: y_values such that each y_values[i] = 2* x_values[i] + 1

<!-- grade_id: cell-d5b1462d5b32f078, points: 0, grade: false, solution: true -->
```python
# Complete the code to - Calculate corresponding y values using x values (above cell) and f(x)

y_values = f(x_values)

y_values
```

Output:
```
array([-5, -1, 11,  5])
```

## Task-3: Linear Combination for Prediction

### Description: Create a function that calculates the linear combination of input data X with weights W and adds the bias b.

#### Input:
#### X: Input data (features).
#### W: Weights.
#### b: Bias term.
#### Output: Predicted output from the linear model, computed as output = np.dot(W, X) + b.

<!-- grade_id: cell-0cbb705c59d3224c, points: 0, grade: false, solution: true -->
```python
def forward_propagation(X, W, b):
   
    # compute the sum of the product of the input and the weights
    output = np.dot(W, X) + b

    # return the predicted price
    return output

```

<!-- grade_id: cell-d473bf061cff83fd, points: 2, grade: true, solution: false -->
```python
X_test = np.array([[1, 2], [3, 4]])
W_test = np.array([[0.5, 1.5]])
b_test = 0
output = forward_propagation(X_test, W_test, b_test)
expected_output = np.array([[5, 7]])
assert np.allclose(output, expected_output)  # Should match dot product without bias
```

## Task-4: e power Function

### Description: Implement a function e power(x) that takes an input x and returns the value of the expression e power(x).

#### Input: An integer or a float x.
#### Output: The result of the expression e power(x).

<!-- grade_id: cell-eec303e7f63227a6, points: 0, grade: false, solution: true -->
```python
def e_power(x):

    return np.exp(x)
```

```python
x = 2.0
e_power(x)
```

Output:
```
7.38905609893065
```

<!-- grade_id: cell-fa82d1f65e767efb, points: 1, grade: true, solution: false -->
```python
assert round(e_power(2.0), 2) == 7.39, "The result should be approximately equal to 7.39"
```

## Task-5: log Function

### Description: Implement a function log f(x) that takes an input x and returns the value of the expression log(x).

#### Input: An integer or a float x.
#### Output: The result of the expression log(x)

<!-- grade_id: cell-a85611e6b7ca7b16, points: 0, grade: false, solution: true -->
```python
def log_f(x):
     
    return np.log(x)
```

```python
x = 2.0
log_f(x)
```

Output:
```
0.6931471805599453
```

<!-- grade_id: cell-4cbb7122d1c8bdee, points: 2, grade: true, solution: false -->
```python
assert round(log_f(2.0), 2) == 0.69, "The result should be approximately equal to 0.69"
```