```python
import torch
import torch.nn as nn
```

# Machine Translation

```python
# Step 1: Tokenizer and Vocabulary Creation (Add more words if needed)
sentence_en = "I love Dhoni ."
# sentence_fr = "J' adore l'IA ."
sentence_te = "నాకు ధోని అంటే ఇష్టం ."

word_map_en = {"<pad>": 0, "I": 1, "love": 2, "Dhoni": 3, ".": 4}
word_map_te = {"<pad>": 0, "నాకు": 1, "ధోని": 2, "అంటే": 3, "ఇష్టం": 4, ".": 5}
```

```python
# Tokenizing sentences
def tokenize(sentence, word_map):
    return torch.tensor([word_map[word] for word in sentence.split()])
```

```python
# Tokenize the input and target sentences
input_tensor = tokenize(sentence_en, word_map_en).unsqueeze(0)  # Shape (1, 4)
target_tensor = tokenize(sentence_te, word_map_te).unsqueeze(0)  # Shape (1, 5)
```

```python
input_tensor.shape, target_tensor.shape
```

Output:
```
(torch.Size([1, 4]), torch.Size([1, 5]))
```

```python
input_tensor, target_tensor
```

Output:
```
(tensor([[1, 2, 3, 4]]), tensor([[1, 2, 3, 4, 5]]))
```

# Task: Uses nn.Embedding to generate trainable word vectors for a given vocabulary of size vocab_size with d_model dimensions

<!-- grade_id: cell-617750cc976f401c, points: 0, grade: false, solution: true -->
```python
class Transformer(nn.Module):
    def __init__(self, vocab_size, d_model):
        super(Transformer, self).__init__()
        ### BEGIN SOLUTION
        self.embedding = nn.Embedding(vocab_size, d_model)
        ### END SOLUTION
        
    def forward(self, src):
        ### BEGIN SOLUTION
        src_embedding = self.embedding(src)
        ### END SOLUTION
       
        return src_embedding
```

<!-- grade_id: cell-c445789abd204ac3, points: 90, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS

# Test Function
def test_transformer():
    vocab_size = 10  # 10 words in vocabulary
    d_model = 4      # Each word is represented by a 4D vector
    model = Transformer(vocab_size, d_model)

    # Test Case 1: Single Word Input
    input_tensor = torch.tensor([2])  # Single index
    output = model(input_tensor)
    assert output.shape == (1, d_model), f"Expected (1, {d_model}), got {output.shape}"
    print("Test 1 Passed: Single word embedding shape is correct.")

    # Test Case 2: Multiple Words (1D tensor)
    input_tensor = torch.tensor([1, 3, 5])
    output = model(input_tensor)
    assert output.shape == (3, d_model), f"Expected (3, {d_model}), got {output.shape}"
    print("Test 2 Passed: Multiple word embeddings shape is correct.")

    # Test Case 3: Batch of Sentences (2D tensor)
    input_tensor = torch.tensor([[1, 2, 3], [4, 5, 6]])  # Batch size = 2, Sequence length = 3
    output = model(input_tensor)
    assert output.shape == (2, 3, d_model), f"Expected (2, 3, {d_model}), got {output.shape}"
    print("Test 3 Passed: Batched input embeddings shape is correct.")
    
# Run Tests
test_transformer()

### END HIDDEN TESTS
```

Output:
```
Test 1 Passed: Single word embedding shape is correct.
Test 2 Passed: Multiple word embeddings shape is correct.
Test 3 Passed: Batched input embeddings shape is correct.
```

```python
#  Initialize Model
vocab_size_en = len(word_map_en)
vocab_size_te = len(word_map_te)
d_model = 8
```

# Task: Creates a Transformer model instance "transformer" that generates trainable word embeddings for a given English vocabulary size (vocab_size_en) with d_model-dimensional vectors.

<!-- grade_id: cell-371d413bf75a514a, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
transformer = Transformer(vocab_size_en, d_model)
### END SOLUTION
```

<!-- grade_id: cell-8628e58db2eb4f8a, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS

# Simple Test Case
vocab_size_en = 10  # Example vocabulary size
d_model = 16  # Embedding dimension

# Initialize the Transformer model
transformer_t = Transformer(vocab_size_en, d_model)

# Check if the model is created successfully
assert isinstance(transformer_t, Transformer), " Test Failed: Model initialization incorrect."
print("Test Passed: Transformer model initialized correctly.")

### END HIDDEN TESTS
```

Output:
```
Test Passed: Transformer model initialized correctly.
```

# Task: Stores the generated word embeddings in "output" by passing input_tensor through the Transformer model.

<!-- grade_id: cell-2f934042bd16d599, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
output = transformer(input_tensor)
### END SOLUTION
```

```python
output
```

Output:
```
tensor([[[ 0.6427,  1.0418, -1.9301, -0.5989, -1.6324,  0.0365, -1.7540,
          -1.1726],
         [ 0.6253, -0.6434, -0.8869,  0.1791, -1.0080,  0.6745,  0.6509,
          -0.9106],
         [ 0.4746, -1.1906,  0.0895, -1.1018, -0.3435,  0.3036,  0.6882,
          -0.5738],
         [-0.2207, -1.3925, -1.1076, -0.5746,  0.2375,  0.3075, -0.2857,
           1.0026]]], grad_fn=<EmbeddingBackward0>)
```

```python

```