# Step 1: Tokenizer and Vocabulary Creation (Add more words if needed)

## 1. Import Necessary Libraries

Import the required PyTorch modules.

<!-- grade_id: cell-93e1006ec1e1082c, points: 0, grade: false, solution: true -->
```python
import torch
import torch.nn as nn
import torch.nn.functional as F
```

You are working on building a simple machine translation model. The first step in the pipeline is to tokenize the given sentences and create vocabulary mappings for both the source and target languages.


```python
# Step 1: Tokenizer and Vocabulary Creation (Add more words if needed)
sentence_en = "I love AI ."
# sentence_fr = "J' adore l'IA ."
sentence_fr = "నాకు ధోని అంటే ఇష్టం ."

word_map_en = {"<pad>": 0, "I": 1, "love": 2, "AI": 3, ".": 4}
word_map_fr = {"<pad>": 0, "నాకు": 1, "ధోని": 2, "అంటే": 3, "ఇష్టం": 4, ".": 5}
```

## 2. Tokenizing sentences

Implement the tokenize function to map each word in the sentence to the corresponding index from the word_map.
Ensure that the function returns a PyTorch tensor containing the word indices.

<!-- grade_id: cell-d7aeab8af7471314, points: 0, grade: false, solution: true -->
```python
# Tokenizing sentences
def tokenize(sentence, word_map):
    return torch.tensor([word_map[word] for word in sentence.split()])
```

## 3.Tokenize the input and target sentences

Tokenize the input (sentence_en) and target (sentence_fr) sentences.
Add the extra dimension to each tokenized tensor using .unsqueeze(0) to make the shape (1, N).

<!-- grade_id: cell-5df70d1e37703f35, points: 0, grade: false, solution: true -->
```python
# Tokenize the input and target sentences
input_tensor = tokenize(sentence_en, word_map_en).unsqueeze(0)  # Shape (1, 6)
target_tensor = tokenize(sentence_fr, word_map_fr).unsqueeze(0)  # Shape (1, 6)
```

# Step 2: Positional Encoding

You are tasked with implementing a Positional Encoding layer in PyTorch, which will be used to inject information about the position of tokens in a sequence. This encoding is crucial for models like the Transformer, which lacks any inherent notion of token position in its architecture.
1.Positional Encoding Initialization:

The PositionalEncoding class will be initialized with the following parameters:
d_model: The dimension of the embeddings (e.g., size of word vectors).
max_len: The maximum sequence length that can be processed (default is 5000).
The positional encoding should be created based on the formula using sine and cosine functions.

2.Encoding Calculation:

You need to calculate the positional encoding matrix using the suitable formula for even indices and odd indices.

This positional encoding matrix will be added to the input sequence in the forward method.

3.Forward Pass:

In the forward method, the input tensor x will be added to the positional encoding corresponding to its sequence length.
The output tensor will be the sum of the input tensor and the positional encoding, which gives the input tokens with position information.
Expected Input and Output:
Given an input sequence of shape (batch_size, seq_len, d_model)

<!-- grade_id: cell-d0ff14153214f8c0, points: 0, grade: false, solution: true -->
```python
class PositionalEncoding(nn.Module):
    def __init__(self, d_model, max_len=5000):
        super(PositionalEncoding, self).__init__()
        self.encoding = torch.zeros(max_len, d_model)
        position = torch.arange(0, max_len, dtype=torch.float).unsqueeze(1)
        div_term = torch.exp(torch.arange(0, d_model, 2).float() * -(torch.log(torch.tensor(10000.0)) / d_model))
        self.encoding[:, 0::2] = torch.sin(position * div_term)
        self.encoding[:, 1::2] = torch.cos(position * div_term)
        self.encoding = self.encoding.unsqueeze(0)

    def forward(self, x):
        return x + self.encoding[:, :x.size(1)]
```

# Test Case 1 : Positional Encoding

<!-- grade_id: cell-70f4cecc8b08269b, points: 100, grade: true, solution: false -->
```python
# Positional Encoding Test Case
import torch
import torch.nn as nn
import numpy as np

def test_positional_encoding():
    d_model = 16  # Smaller model for testing
    max_len = 10  # Small sequence length
    pos_enc = PositionalEncoding(d_model, max_len)

    # Create a mock input tensor
    x = torch.zeros((1, max_len, d_model))  # Batch size = 1, seq_len = 10, d_model = 16

    # Forward pass
    output = pos_enc(x)

    # Check if the output shape is correct
    assert output.shape == x.shape, "Output shape mismatch"

    # Check if the encoding changes the input
    assert not torch.allclose(output, x), "Positional encoding is not modifying the input"

    # Check for consistency in encoding values
    encoding1 = pos_enc.encoding[:, :max_len].detach().numpy()
    encoding2 = pos_enc.encoding[:, :max_len].detach().numpy()
    assert np.allclose(encoding1, encoding2), "Positional encoding is not consistent across calls"

    print("Test Case 1 :All positional encoding tests passed!")

# Run test case
test_positional_encoding()

```

Output:
```
Test Case 1 :All positional encoding tests passed!
```

# Step 3: Multi-Head Attention

Your task is to implement the MultiHeadAttention class, a key component of the Transformer model. This module enables the model to focus on different parts of the input sequence simultaneously using multiple attention heads.

You will implement the following functionalities:
1.Initialization:

    Define the parameters d_model (input and output dimensionality) and num_heads (number of attention heads).
    Compute d_k and d_v, the dimensionality of each attention head.
    Define linear transformations for computing query (self.query), key (self.key), and value (self.value).
    Define the final fully connected layer (self.fc) to project the concatenated outputs back to d_model.
2.Forward Pass:

    Compute query (q), key (k), and value (v) vectors by applying the respective linear transformations.
    Reshape q, k, and v to split them across multiple heads, ensuring each head gets d_k and d_v dimensions.

    Compute scaled dot-product attention using the formula:

     Apply a mask if provided (mask), ensuring masked positions do not influence attention.

    Compute attention weights (attn_weights) using the Softmax function.

    Apply the attention weights to the value (v) tensor, producing attention_output.

    Reshape and concatenate all attention heads to get back to the original d_model dimension.

    Pass the result through the final fully connected layer (self.fc) to obtain the final attention output.

<!-- grade_id: cell-29bab5f4ee918a17, points: 0, grade: false, solution: true -->
```python
# Step 3: Multi-Head Attention
class MultiHeadAttention(nn.Module):
    def __init__(self, d_model, num_heads):
        super(MultiHeadAttention, self).__init__()
        self.num_heads = num_heads
        self.d_model = d_model
        self.d_k = d_model // num_heads
        self.d_v = d_model // num_heads

        self.query = nn.Linear(d_model, d_model)
        self.key = nn.Linear(d_model, d_model)
        self.value = nn.Linear(d_model, d_model)

        self.fc = nn.Linear(d_model, d_model)

    def forward(self, x, mask=None):
        batch_size, seq_len, _ = x.size()

        q = self.query(x).view(batch_size, seq_len, self.num_heads, self.d_k).transpose(1, 2)
        k = self.key(x).view(batch_size, seq_len, self.num_heads, self.d_k).transpose(1, 2)
        v = self.value(x).view(batch_size, seq_len, self.num_heads, self.d_v).transpose(1, 2)

        attn_scores = torch.matmul(q, k.transpose(-2, -1)) / torch.sqrt(torch.tensor(self.d_k, dtype=torch.float))
        if mask is not None:
            attn_scores = attn_scores.masked_fill(mask == 0, float('-inf'))  # Ensure mask matches input length
        attn_weights = F.softmax(attn_scores, dim=-1)

        attention_output = torch.matmul(attn_weights, v).transpose(1, 2).contiguous().view(batch_size, seq_len, self.d_model)
        output = self.fc(attention_output)

        return output
```

# Test Case 2 Check multi_head_attention

<!-- grade_id: cell-190e4c0b25291225, points: 100, grade: true, solution: false -->
```python
import torch
import torch.nn as nn
import torch.nn.functional as F
import numpy as np

def test_multi_head_attention():
    d_model = 16  # Small model for testing
    num_heads = 4  # Number of heads
    seq_len = 5  # Sequence length
    batch_size = 2  # Batch size

    mha = MultiHeadAttention(d_model, num_heads)

    # Create a random input tensor
    x = torch.rand((batch_size, seq_len, d_model))

    # Create a random mask
    mask = torch.randint(0, 2, (batch_size, num_heads, seq_len, seq_len))

    # Forward pass
    output = mha(x, mask)

    # Test 1: Check multi-head splits
    q = mha.query(x).view(batch_size, seq_len, num_heads, d_model // num_heads).transpose(1, 2)
    assert q.shape == (batch_size, num_heads, seq_len, d_model // num_heads), "Query split is incorrect"

    # Test 2: Check softmax attention weights sum to 1
    attn_scores = torch.matmul(q, q.transpose(-2, -1)) / torch.sqrt(torch.tensor(d_model // num_heads, dtype=torch.float))
    attn_weights = F.softmax(attn_scores, dim=-1)
    attn_sums = attn_weights.sum(dim=-1)
    assert torch.allclose(attn_sums, torch.ones_like(attn_sums), atol=1e-5), "Attention weights do not sum to 1"

    print("Test Case 2:All Multi-Head Attention tests passed!")

# Run test case
test_multi_head_attention()

```

Output:
```
Test Case 2:All Multi-Head Attention tests passed!
```

# Step 4: Feed-Forward Network

You are tasked with implementing a Feed-Forward Network (FFN), a critical part of the Transformer architecture. The FFN processes each token's representation independently, consisting of two fully connected layers. The first layer transforms the input features into a larger space (d_ff), and the second layer projects it back to the original feature dimension (d_model). A ReLU activation function is applied between these two layers.

1.Initialization:
    The FeedForward class should be initialized with:
    d_model: The dimensionality of the input and output feature vectors.
    d_ff: The dimensionality of the hidden layer (default value is 512).

2.Forward Pass:
    The forward pass will apply two fully connected layers:
    First, the input tensor x (of shape [batch_size, seq_len, d_model]) will be passed through a linear layer (fc1) that projects it to the higher dimension d_ff.
    Then, a ReLU activation function is applied to the output of fc1.
    Finally, the result will be passed through a second linear layer (fc2), which projects the output back to the original dimension d_model.
3.Expected Input and Output:
    Given an input tensor x of shape [batch_size, seq_len, d_model], the output will be the same shape [batch_size, seq_len, d_model], but with transformed features.

<!-- grade_id: cell-a2ecdf8a3b200e2a, points: 0, grade: false, solution: true -->
```python
# Step 4: Feed-Forward Network
class FeedForward(nn.Module):
    def __init__(self, d_model, d_ff=512):
        super(FeedForward, self).__init__()
        self.fc1 = nn.Linear(d_model, d_ff)
        self.fc2 = nn.Linear(d_ff, d_model)

    def forward(self, x):
        return self.fc2(F.relu(self.fc1(x)))
```

# Step 5: Encoder Layer

You are tasked with implementing an Encoder Layer for a Transformer model. The encoder layer consists of two main components:

Multi-Head Attention: This mechanism allows the model to focus on different parts of the input sequence simultaneously.
Feed-Forward Network (FFN): This component applies two fully connected layers to each position independently.
In addition, each of these components is followed by Layer Normalization and Residual Connections to stabilize training.

Task Breakdown:
1.Initialization:
    The EncoderLayer class should be initialized with:
    d_model: The dimensionality of the input and output features.
    num_heads: The number of attention heads in the multi-head attention mechanism.
    d_ff: The dimensionality of the hidden layer in the feed-forward network (default is 512).
    The following components should be initialized:
    multihead_attn: A MultiHeadAttention layer.
    feedforward: A FeedForward layer.norm1, norm2: Two Layer Normalization layers.

2.Forward Pass:
The input tensor x is first passed through the multi-head attention mechanism (multihead_attn). The output is added to the original input x (residual connection), and then Layer Normalization (norm1) is applied.
The resulting tensor is then passed through the feed-forward network (feedforward), and again the residual connection is added followed by Layer Normalization (norm2).
The output of this forward pass will be the transformed tensor with attention and feed-forward operations applied.



<!-- grade_id: cell-ac021edf972a02a8, points: 0, grade: false, solution: true -->
```python
# Step 5: Encoder Layer
class EncoderLayer(nn.Module):
    def __init__(self, d_model, num_heads, d_ff=512):
        super(EncoderLayer, self).__init__()
        self.multihead_attn = MultiHeadAttention(d_model, num_heads)
        self.feedforward = FeedForward(d_model, d_ff)
        self.norm1 = nn.LayerNorm(d_model)
        self.norm2 = nn.LayerNorm(d_model)

    def forward(self, x, mask=None):
        attn_output = self.multihead_attn(x, mask)
        x = self.norm1(x + attn_output)
        ff_output = self.feedforward(x)
        x = self.norm2(x + ff_output)
        return x
```

# Test Case 3: Check Encoder Layer

<!-- grade_id: cell-975f0ea7583cd5ff, points: 100, grade: true, solution: false -->
```python
import torch
import torch.nn as nn

def test_encoder_layer():
    d_model = 16  # Model dimension
    num_heads = 4  # Number of attention heads
    seq_len = 5  # Sequence length
    batch_size = 2  # Batch size

    encoder_layer = EncoderLayer(d_model, num_heads)

    # Create a random input tensor
    x = torch.rand((batch_size, seq_len, d_model))

    # Create a random mask
    mask = torch.randint(0, 2, (batch_size, num_heads, seq_len, seq_len))

    # Forward pass
    output = encoder_layer(x, mask)

    # Test 1: Check that MultiHeadAttention modifies input
    attn_output = encoder_layer.multihead_attn(x, mask)
    assert not torch.equal(attn_output, x), "MultiHeadAttention does not modify input"

    # Test 2: Check that MultiHeadAttention modifies input
    x = torch.rand(10, 20, d_model)
    output = encoder_layer(x)
    assert torch.all(output.std(dim=-1) > 0), "LayerNorm not applied correctly!"

    print("Test Case 3:EncoderLayer tests passed!")

# Run test case
test_encoder_layer()

```

Output:
```
Test Case 3:EncoderLayer tests passed!
```

# Step 6: Decoder Layer

You need to implement the DecoderLayer class with the following requirements:

1.Initialization (__init__ method):
    Accepts d_model (hidden dimension), num_heads (number of attention heads), and d_ff (hidden size for the feed-forward network, default is 512).
Initializes:
    multihead_attn1: First multi-head self-attention layer (masked self-attention).
    multihead_attn2: Second multi-head attention layer (cross-attention with encoder output).
    feedforward: Fully connected feed-forward network.
    norm1, norm2, and norm3: Three Layer Normalization layers.

Forward Pass (forward method):
Accepts the following inputs:
    x: Target sequence input tensor (shape: [batch_size, tgt_seq_len, d_model]).
    encoder_output: Output from the encoder (shape: [batch_size, src_seq_len, d_model]).
    tgt_mask: Mask to prevent looking at future tokens in the target sequence.
    src_mask: Mask to ignore padding in the encoder’s output.
    Applies:
    1.Masked Multi-Head Self-Attention (multihead_attn1):
    Ensures the model attends only to previous words in the target sequence.
    Output(attn_output1) is added to x (residual connection) and normalized with norm1.
    2.Multi-Head Attention (multihead_attn2):
    Attends to the encoder’s output to retrieve relevant information.
    Output(attn_output2) is added to x (residual connection) and normalized with norm2.
    3.Feed-Forward Network (feedforward):
    Processes the result through fully connected layers.
    Output(ff_output) is added to x (residual connection) and normalized with norm3.

<!-- grade_id: cell-a089f6447080b612, points: 0, grade: false, solution: true -->
```python
# Step 6: Decoder Layer
class DecoderLayer(nn.Module):
    def __init__(self, d_model, num_heads, d_ff=512):
        super(DecoderLayer, self).__init__()
        self.multihead_attn1 = MultiHeadAttention(d_model, num_heads)
        self.multihead_attn2 = MultiHeadAttention(d_model, num_heads)
        self.feedforward = FeedForward(d_model, d_ff)
        self.norm1 = nn.LayerNorm(d_model)
        self.norm2 = nn.LayerNorm(d_model)
        self.norm3 = nn.LayerNorm(d_model)

    def forward(self, x, encoder_output, tgt_mask=None, src_mask=None):
        attn_output1 = self.multihead_attn1(x, tgt_mask)
        x = self.norm1(x + attn_output1)
        attn_output2 = self.multihead_attn2(x, src_mask)
        x = self.norm2(x + attn_output2)
        ff_output = self.feedforward(x)
        x = self.norm3(x + ff_output)
        return x
```

# Test Case 4 : Decoder Layer

<!-- grade_id: cell-a3f9189dca98c048, points: 100, grade: true, solution: false -->
```python

def test_decoder_layer():
    d_model = 512
    num_heads = 8
    seq_length = 10
    batch_size = 4

    # Create a random tensor as input
    x = torch.rand(batch_size, seq_length, d_model)
    encoder_output = torch.rand(batch_size, seq_length, d_model)

    # Initialize the DecoderLayer
    decoder_layer = DecoderLayer(d_model, num_heads)
    assert hasattr(decoder_layer, "multihead_attn2"), "DecoderLayer must have a cross-attention layer"
    assert hasattr(decoder_layer, "feedforward"), "DecoderLayer must have a feed-forward network"
    # Forward pass
    output = decoder_layer(x, encoder_output)

    # Assertions
    assert output.shape == x.shape, "Output shape mismatch"
    assert not torch.isnan(output).any(), "Output contains NaN values"

# Run the test
test_decoder_layer()
print("Test Case 4 Passed")
```

Output:
```
Test Case 4 Passed
```

# Step 7: Transformer Model

Your task is to implement a Transformer model from scratch in PyTorch, which can process an input sequence (src) and generate an output sequence (tgt). The Transformer consists of an encoder-decoder architecture where:

The Encoder processes the input (src) and generates a contextualized representation.
The Decoder takes this representation along with the target (tgt) to generate the output.
You need to implement the Transformer class using the provided EncoderLayer and DecoderLayer modules.


1.Initialize the following layers in the __init__ method:

    An embedding layer (self.embedding) to convert input token indices into vector representations.
    A positional encoding layer (self.pos_encoder) to provide positional information.
    A stack of encoder layers (self.encoder_layers) using nn.ModuleList.
    A stack of decoder layers (self.decoder_layers) using nn.ModuleList.
    A fully connected output layer (self.fc_out) that maps the decoder output to the vocabulary space.

2.Implement the forward method:
    Embed the src and tgt sequences and apply positional encoding.
    Pass src through the encoder layers sequentially.
    Pass tgt through the decoder layers sequentially, using the encoder output.
    Apply the final linear layer to get the output logits

<!-- grade_id: cell-9dde5c6dd1daa905, points: 0, grade: false, solution: true -->
```python
# Step 7: Transformer Model
class Transformer(nn.Module):
    def __init__(self, vocab_size, d_model, num_heads, num_encoder_layers, num_decoder_layers, max_len=5000):
        super(Transformer, self).__init__()
        self.embedding = nn.Embedding(vocab_size, d_model)
        self.pos_encoder = PositionalEncoding(d_model, max_len)
        self.encoder_layers = nn.ModuleList([EncoderLayer(d_model, num_heads) for _ in range(num_encoder_layers)])
        self.decoder_layers = nn.ModuleList([DecoderLayer(d_model, num_heads) for _ in range(num_decoder_layers)])
        self.fc_out = nn.Linear(d_model, vocab_size)

    def forward(self, src, tgt, tgt_mask=None):
        src = self.pos_encoder(self.embedding(src))
        tgt = self.pos_encoder(self.embedding(tgt))

        encoder_output = src
        for layer in self.encoder_layers:
            encoder_output = layer(encoder_output)

        decoder_output = tgt
        for layer in self.decoder_layers:
            decoder_output = layer(decoder_output, encoder_output, tgt_mask=tgt_mask)

        output = self.fc_out(decoder_output)
        return output
```

# Step 8: Prediction Module

Your task is to implement the translate function, which takes an English sentence as input and returns the translated French sentence using a Transformer model.

This function should perform the following steps:
1.Convert the input sentence into tokens using word_map_en.
2.Create a target tensor initialized with zeros.
3.Use the Transformer model to generate a translated sentence.
4.Apply Softmax to compute token probabilities.
5.Use Argmax to determine the most probable token at each position.
6.Convert token indices back to words using word_map_fr.

<!-- grade_id: cell-b78400dfce5aba3f, points: 0, grade: false, solution: true -->
```python
# Step 8: Prediction Module
def translate(input_sentence, word_map_en, word_map_fr, transformer):
    # Tokenize input sentence
    input_tensor = tokenize(input_sentence, word_map_en).unsqueeze(0)  # Shape (1, seq_len)

    # Generate a mask for the target sentence
    tgt_mask = torch.tril(torch.ones((input_tensor.size(1), input_tensor.size(1)))).unsqueeze(0).unsqueeze(0)  # Lower triangular mask

    # Initialize a tensor for the target sentence
    target_tensor = torch.zeros((1, input_tensor.size(1)), dtype=torch.long)

    # Predict the output sentence (translation)
    output = transformer(input_tensor, target_tensor, tgt_mask)

    # Apply Softmax to get probabilities
    softmax_output = F.softmax(output, dim=-1)

    # Get predicted token indices (Argmax)
    predicted_tokens = torch.argmax(softmax_output, dim=-1)

    # Convert predicted tokens back to words
    reverse_word_map_fr = {v: k for k, v in word_map_fr.items()}
    translated_sentence = [reverse_word_map_fr[token.item()] for token in predicted_tokens[0] if token != 0]

    return " ".join(translated_sentence)
```

# Step 9: Initialize Model

```python
# Step 9: Initialize Model
vocab_size_en = len(word_map_en)
vocab_size_fr = len(word_map_fr)
d_model = 128
num_heads = 8
num_encoder_layers = 2
num_decoder_layers = 2

transformer = Transformer(vocab_size_en, d_model, num_heads, num_encoder_layers, num_decoder_layers)
```

# Step 10: Take Input Sentence and Generate Translation

```python
# Step 10: Take Input Sentence and Generate Translation
input_sentence = "I love AI ."  # Example input sentence
predicted_sentence = translate(input_sentence, word_map_en, word_map_fr, transformer)
print("Input Sentence:", input_sentence)
print("Predicted Translation:", predicted_sentence)
```

Output:
```
Input Sentence: I love AI .
Predicted Translation: ధోని ధోని ధోని ధోని
```

```python

```

```python

```